Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Note that this Pre-class Work is estimated to take **63 minutes**.

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Andriy Kashyrskyy"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work - Computational Applications of Randomized Algorithms

## Part A. Median Heap (watch this [video explanation](https://www.youtube.com/watch?v=756_8C9YBZQ&list=PLF_a-qBXTGFektoI6JUOTRL36JlvD04BR&index=5&t=0s) or read this [description](https://stackoverflow.com/a/15319593/7946759))

Throughout this pre-class work, please use the following definition of median: the median of a list of numbers is the one in the middle of the list when the list is ordered. When such the middle element can’t be determined (i.e., in a list of even length), the average of the two middle elements is the median. For example, 5 is the median of [-1,2,4,5,8,10,12], and (5+7)/2=6 is the median of [1,2,3,5,7,8,10,11].

Using the idea from Lesson Heaps and Heapsort, we can use a pair of heaps to create a data structure which allows fast access to the median. Use the heapq module in python to create both a max-heap and a min-heap. Note that by default, the heapq module in python only creates min-heaps, but if we multiply elements by -1 when we store them, then we can also create max-heaps.

## Question 1 [time estimate: 5 minutes]
Write a function `median(minh, maxh)`. It must return the median element.


<!-- grade_id: cell-43ce866dd1dc2d5b, points: 0, grade: false, solution: true -->
```python
def median(minh, maxh):
    med = 0
    
    if len(minh) > len(maxh):
        med = minh[0]
    if len(minh) < len(maxh):
        med = maxh[0]
    if len(minh) == len(maxh) and len(minh)!=0:
        med = (minh[0]+maxh[0])/2
        
    return med
```

## Question 2 [time estimate: 10 minutes]
Write a function `add_to_median_heap(minh, maxh, elem)`. It must accept a min heap, a max heap, and an element to add.


<!-- grade_id: cell-50cb08de70712ee7, points: 0, grade: false, solution: true -->
```python
import heapq

def add_to_median_heap(minh, maxh, elem):
    
    med = median(minh, maxh)
    
    if elem > med:
        heapq.heappush(minh, elem)
    else:
        heapq.heappush(maxh, -elem)
    
    if len(minh) == len(maxh) + 2:
        i = heapq.heappop(minh)
        heapq.heappush(maxh, -i)
    elif len(maxh) == len(minh) + 2:
        i = -heapq.heappop(maxh)
        heapq.heappush(minh, i)
```

<!-- grade_id: cell-673faa69bf330eda, points: 0, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

## Question 3 [time estimate: 1 minute]

Uncomment and run the testing code given below to test your functions. It should print out numbers ranging from 1 to 50, in that order.

```python
minh = []
maxh = []
for a in range(1,100,2):
    add_to_median_heap(minh, maxh, a)
    print(median(minh, maxh))  
```

Output:
```
1
1.0
3
1.0
5
1.0
7
1.0
9
1.0
11
1.0
13
1.0
15
1.0
17
1.0
19
1.0
21
1.0
23
1.0
25
1.0
27
1.0
29
1.0
31
1.0
33
1.0
35
1.0
37
1.0
39
1.0
41
1.0
43
1.0
45
1.0
47
1.0
49
1.0
```

## Question 4 [time estimate: 3 minutes]
What’s the worst case complexity to build a median heap using `add_to_median_heap`?

<!-- grade_id: cell-c8e33ebdc8aaf328, points: 0, grade: true, solution: true -->
Heappop and heappush functions are both taking O(logn) time complexities, thus total will be O(logn + logn) = O(2logn), or just O(logn).

## Question 5 [time estimate: 3 minutes] 
What’s the worst case complexity of `median`?

<!-- grade_id: cell-c061920faa337890, points: 0, grade: true, solution: true -->
It should just be O(1).

## Question 6 [time estimate: 4 minutes]

How does this way of finding the median compare with the vanilla way of sorting the list and pick the middle element? Use arguments based on efficiency or clearity of the respective algorithms.

<!-- grade_id: cell-1f31afecda9730a7, points: 0, grade: true, solution: true -->
I believe this approach is more efficient, as sorting and picking a median for some algorithms can take O(n^2), while our approach of calculating a median takes O(nlogn).

## [Optional] Question 7 [time estimate: 15 minutes]

Is it possible to extend this idea to any percentile? If it is, then write code to do so. If it’s not possible, prove why it is not possible.

<!-- grade_id: cell-4d0c197e6d484338, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
# raise NotImplementedError()
```

## Part B. Quickselect

Quicksort can be modified to find the $k$-th smallest element in an unordered list. This is known as quickselect. It does this by choosing a partition (as in quicksort). Once the list has been partitioned then we know how many elements lie to the left and to the right of the partition. This allows us to recursively call quickselect on the correct sublist.

## Question 1 [time estimate: 10 minutes]

Write a function `qselect(lst, k)`, which takes both a list and an index $k$. The function must then return the $k$-th smallest item in the list.


<!-- grade_id: cell-fb9c2bc9d5b4bc92, points: 0, grade: false, solution: true -->
```python
## My idea is to use the partition and quick_sort from 
## previous Pre-Class work. I could not solve it, however.

## Using code from my PCW submissiong for Deterministic Quicksort Session 7.2.
def partition(A,p,r):
    """
    Assume r<len(A) and p>=0
    
    Note: This function picks the last element as pivot, 
    places it in the correct position in a sorted array
    (all smaller elements before it to the left, 
    all bigger - after, to the right).
    """
    x = A[r] ## pivot
    i = p - 1 ## index of a smaller element
    
    ## if current elementt is smaller than or equal to pivot
    for j in range(p, r):
        if A[j] <= x:
            ## increment index of smaller element
            i = i + 1
            A[i], A[j] = A[j], A[i] 
            
    A[i+1], A[r] = A[r], A[i+1]
    return (i+1)

def quick_sort(A,p,r):
    if len(A) == 1: ## simple base case condition if array is of length 1
        return A ## returning initial array of 1, as it is already sorted
    if p < r:
        ## q is a partitioning index
        q = partition(A,p,r) ## 
        quick_sort(A, p, q-1) ## calling quick_sort on left side
        quick_sort(A, q+1, r) ## calling quick sort on right side
        return A ## returning sorted array
```

```python
def qselect(lst,k):
    
```

<!-- grade_id: cell-85a25ed6a1f62a66, points: 0, grade: true, solution: false -->
```python
import random
random.seed(123) # introducing a seed for reproducibility purposes
lst1 = list(range(100))
random.shuffle(lst1)
lst2 = []
for a in range(100):
    lst2.append(qselect(lst1, a))
assert(lst2 == list(range(100)))
```

Output:
```
<ipython-input-43-253afcee666d> in <module>
      5 lst2 = []
      6 for a in range(100):
----> 7     lst2.append(qselect(lst1, a))
      8 assert(lst2 == list(range(100)))

<ipython-input-42-9d4d243ddd21> in qselect(lst, k)
      3     r = len(lst)-1
      4 
----> 5     quick_sort(lst,p,r,k)
      6     return(lst[k])

TypeError: quick_sort() takes 3 positional arguments but 4 were given
```

## Question 2 [time estimate: 1 minute]

Uncomment and run the testing code given below to test your function. It should print out integers from 0 to 99.

```python
# import random
# random.seed(123) # introducing a seed for reproducibility purposes
# lst = list(range(100))
# random.shuffle(lst)
# for a in range(100):
#     print(qselect(lst, a))
```

## Question 3 [time estimate: 5 minutes]

Write down the recurrence relation for your code.

<!-- grade_id: cell-24144607b3ffade3, points: 0, grade: true, solution: true -->
I will figure out how to code qselect(), and will then finish the remaining Questions 3, 4, 5.

## Question 4 [time estimate: 3 minutes]

Solve the recurrence relation for quickselect in the best case.


<!-- grade_id: cell-c125f4742933e820, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 5 [time estimate: 3 minutes]

Solve the recurrence relation for quickselect in the worst case.

<!-- grade_id: cell-852c3a71d8c2d7c1, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE