Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Note that this Pre-class Work is estimated to take **32 minutes**.

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Andriy Kashyrskyy"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work - Maximum subarray problem—part II

## Question 1 [time estimate: 2 minutes]
Paste in your Python implementation of the maximum subarray from the previous class in the cell below and use that to find out the value of the maximum subarray of this array: `A = [-2, -3, 4, -1, -2, 1]`

<!-- grade_id: cell-e69c8bbcd40ca242, points: 0, grade: true, solution: true -->
```python
def bruteforce_max_subarray(A):
    """
    Implements brute-force maximum subarray finding.
    
    Parameters
    ----------
    A : list
        a NON-EMPTY list of floats
    
    Returns
    -------
    tuple
        - the start index of the max subarray
        - the end index of the max subarray
        - the value of the maximum subarray
        
    """
    ## we create empty lists for indeces of starting array index and ending array index
    ## we initially set maximum to 0 
    ## maximum, index_1, index_2 will be getting updated while the code below runs
    maximum = 0
    index_1 = []
    index_2 = []
    ## as we iterate through each sub-array, we will calculate each of it's sum
    ## current_maximum will store the maximum of these sums seen so far
    ## we can also use range(len(A)-1)
    for a in range(0,len(A)):
        for b in range(0,len(A)):
            if A[a:b]:
                ## compute the sum, save to current max value
                current_maximum = sum(A[a:b])
                ## if the maximum is less than current maximum, append a value and b value to indeces lists
                ## update maximum with a higher, current maximum
                if current_maximum > maximum:
                    maximum = current_maximum
                    index_1.append(a)
                    index_2.append(b-1)
   ## we output last indeces in the arrays of indeces and a maximum sum
    return index_1[-1], index_2[-1], maximum

A = [-2, -3, 4, -1, -2, 1]
bruteforce_max_subarray(A)
```

Output:
```
(2, 2, 4)
```

## Question 2 [time estimate: 5 minutes]
Now, your friend Joe comes and appends a single extra number at the end of the array, which becomes:  `B = [-2, -3, 4, -1, -2, 1, 8]`. With the output from ```max_subarray(A)```, do you need to rerun ```max_subarray()``` for the new array again to find the subarray we need? Explain why, or why not, in as much detail as possible, using, if necessary, a specific example to illustrate your reasoning. The subsequent questions will help you figure out an efficient algorithmic strategy to address the last question, but make sure to write your explanation above first, before answering the remaining questions.


<!-- grade_id: cell-d4b7cd845c816263, points: 0, grade: true, solution: true -->
If we do run it again with the new array, we can expect to have a larger run time (as we get a larger input), which results in a bigger number of steps performed (in a brute-force it would be n(n-1)/2; so if n(n-1)/2 is what we have originally, the next new can be calculated as (n+1)(n)/2 is we just updated n by 1).

With the added *8* to the right, we can see that the new maximum subarray will be *10* (as values of 4+(-1)+(-2)+1+8). Regarding this question, to figure out such problem I would just calculate whether the sum of each array including 8 (starting from the end) minus each sum calculated before, and the biggest sum we'll end up with will correspond (as a positive number) to a bigger max subarray then calculated before we added an extra number.

## Question 3 [time estimate: 5 minutes]

**Determine if the following statement is True or False and explain your answer.**
If the maximum subarray of the array A is different than the maximum subarray of the array B (questions 1 and 2), the new maximum-subarray doesn’t need to include 8 (i.e., the newly appended element). 


<!-- grade_id: cell-9377964a8756f13b, points: 0, grade: true, solution: true -->
*False if max_subarray(B)>maxsubarray(A)*. We have to define what "different" means in this context, and we can define this as having maximum subarray of the array A either smaller or bigger then a maximum subarray of the array B. In the case of it being smaller, the new maximum-subarray HAS to include 8, as that makes a difference in its algorithmic computation (we add one additional number). If it is bigger, that means adding 8 to the array did not result in having a bigger value for the maximum subarray. However, as we have seen before, adding of 8 results in a new subarray (4,-1,-2,1,8) with a sum of 10, which is bigger than A's maximum subarray of 4.

## Question 4 [time estimate: 10 minutes]
Complete the Python function `incremental_max_subarray(x, mx)` in the cell below.

This [video](https://www.youtube.com/watch?v=AAgErqQmwmA&list=PLF_a-qBXTGFektoI6JUOTRL36JlvD04BR&index=4&t=0s)  might be helpful to understand the `incremental_max_subarray` problem.

<!-- grade_id: cell-0230e459f3d701f8, points: 0, grade: false, solution: true -->
```python
def incremental_max_subarray(x, mx):
    """ 
    Parameters
    ----------
    x : list
        A NON-EMPTY list of numbers (e.g., the 
        array B in the first two questions above)
        
        If x has 1 element: returns the value of 
        the element regardless of the value of mx
        
    mx 
        The maximum subarray of x excluding its 
        ast element (i.e., compute the maximum 
        subarray of the input array x considering 
        only its first len(x) - 1 elements)
    
    Returns
    -------
    int
        The maximum subarray of the array x
    
    Example
    -------
    Using the array B in question 2, the result 
    of incremental_max_subarray(B, 4) is 10: 
    
        10 = 8 + 1 - 2 -1 + 4 
    
    
    """
    ## specifying the base case condition
    if  len(x) == 1:
        return x[0]
    
    minor_sum = 0
    final_sum = -float("inf")
    
    ## making a for loop to iterate through the array
    for a in reversed(range(-len(x),0)):
        minor_sum = minor_sum + x[a] #adding the ath element
        if minor_sum > final_sum: #if subsum we find is bigger than global sum
            final_sum = minor_sum # updating the final (global sum), if meets conditions
            
    return max(mx, final_sum) #returning mx, final_sum

B = [-2, -3, 4, -1, -2, 1, 8]

## finding the maximum subarray first
BB = B[:-1]
mxx = int(bruteforce_max_subarray(BB)[2])
## calling on the function with a new list
incremental_max_subarray(B, mxx)
```

Output:
```
10
```

<!-- grade_id: cell-9abf19e56620ffa3, points: 1, grade: true, solution: false -->
```python
B = [-2, -3, 4, -1, -2, 1, 8]
assert(incremental_max_subarray(B, 4) == 10)
assert(incremental_max_subarray(B[:1], 0) == B[0])
```

## Question 5

### Part 1 [time estimate: 5 minutes]

Now use `incremental_max_subarray(x, mx)` iteratively, starting from the first element, to calculate the maximum subarray of a given list. Come up with a few test cases to test your code.

<!-- grade_id: cell-fd96d4ccccd99fe6, points: 0, grade: false, solution: true -->
```python
def max_subarray(A):
    """
    Using `incremental_max_subarray` iteratively on A 
    to produce the value of the maximum subarray of A.
    
    Parameters
    ----------
    A : list
        A NON-EMPTY list of floats
        
    Returns
    -------
    float
        The sum of the maximum subarray of A 
    
    """
    #base-case
    #if  len(A) == 1:
    sum_global = A[0]
    
    #running a for loop to add new values and comppute the new (global) sum
    for a in range(len(A)+1):
        sum_global = incremental_max_subarray(A[:a], sum_global)    
    return sum_global

A = [-2, -3, 4, -1, -2, 1]
B = [-2, -3, 4, -1, -2, 1, 8]
assert(max_subarray(A) == 4)
assert(max_subarray(B) == 10)
```

<!-- grade_id: cell-669ad779766aa2c3, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

### Part 2 [time estimate: 5 minutes]
Is this more efficient than the divide-and-conquer approach? Explain.

<!-- grade_id: cell-cd8f0b130a7136db, points: 0, grade: true, solution: true -->
Its time complexity is O(n^2) (1 for loop within max_subarray() and 1 for loop within incremental_max_subarray) in comparison to that of divide and conquer approach O(nlogn). Thus it is less efficient (with the larger input size of n divide-and-conquer performs better, that is graphically presented by the plot of nlogn gettung lowe values of T(nlogn) on y-axis, in comparison to the same inputs for the plot of T(n^2).