Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Note that this Pre-class Work is estimated to take **37 minutes**.

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Andriy Kashyrskyy"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work - Asymptotic notations; maximum-subarray problem

## Question 1 (Exercise 3.1-3 of Cormen, et al. ) [time estimate: 5 minutes]
Explain why the statement, "The running time of algorithm A is at least $O(n^2)$," is true but meaningless.


<!-- grade_id: cell-437d61c1420d4f99, points: 0, grade: true, solution: true -->
(i) Though true, it is meaningless to say since for a given n the actual running time will vary, depending on the particular input of size n.

(ii) While the statement is partially true, it only covers the upper bound (with O-notation, worst case) of the algorithm's asymptotic analysis, excluding the running time under Theta-, and Omega-notation (related to lower bound (best case), and average case respectively). 

## Question 2 (Exercise 3.1-4 of Cormen, et al. ) [time estimate: 5 minutes]

Is $2^{n+1}=O(2^n)$? Is $2^{2n}=O(2^n)$? Why?

<!-- grade_id: cell-6a0c4ee19dadfddf, points: 0, grade: true, solution: true -->
I believe the first equation is valid, while the second is not because the powers should be the same for the bound to be asymptotically tight; in the second example the bound isn't asymptotically tight as "2n" is not equal to "n". In the first case we can simplify 2^(n+1) as 2(2^(n)), and it is going to be equal to 2^n (since coefficient of 2 isn't considered, it is assumed to have no impact on scaling of this polynomial).

## Question 3 [time estimate: 20 minutes]
Write a function in Python that solves the maximum-subarray problem using a brute-force approach. Your Python function must:
* Take as Input an array/list  of numbers
* Produce the following Output: 
    * the start and end indices of the subarray containing the maximum sum.
    * value of the maximum subarray (float)


<!-- grade_id: cell-527e6532d6992fd0, points: 0, grade: false, solution: true -->
```python
def bruteforce_max_subarray(A):
    """
    Implements brute-force maximum subarray finding.
    
    Parameters
    ----------
    A : list
        a NON-EMPTY list of floats
    
    Returns
    -------
    tuple
        - the start index of the max subarray
        - the end index of the max subarray
        - the value of the maximum subarray
        
    """
    ## we create empty lists for indeces of starting array index and ending array index
    ## we initially set maximum to 0 
    ## maximum, index_1, index_2 will be getting updated while the code below runs
    maximum = 0
    index_1 = []
    index_2 = []
    ## as we iterate through each sub-array, we will calculate each of it's sum
    ## current_maximum will store the maximum of these sums seen so far
    ## we can also use range(len(A)-1)
    for a in range(0,len(A)):
        for b in range(0,len(A)):
            if A[a:b]:
                ## compute the sum, save to current max value
                current_maximum = sum(A[a:b])
                ## if the maximum is less than current maximum, append a value and b value to indeces lists
                ## update maximum with a higher, current maximum
                if current_maximum > maximum:
                    maximum = current_maximum
                    index_1.append(a)
                    index_2.append(b-1)
   ## we output last indeces in the arrays of indeces and a maximum sum
    return index_1[-1], index_2[-1], maximum
```

<!-- grade_id: cell-a28a56466c9537ff, points: 1, grade: true, solution: false -->
```python
assert(bruteforce_max_subarray([-2,1,-1,2,-5]) == (1, 3, 2))
assert(bruteforce_max_subarray([-2, -5, 6, -2, -3, 1, 5, -6]) == (2, 6, 7))
```

## Question 4 [time estimate: 2 minutes]
Test your Python maximum-subarray function using the following input list (from Figure 4.3 of Cormen et al.):  
`A = [13, -3, -25, 20, -3, -16, -23, 18, 20, -7, 12, -5, -22, 15, -4, 7] `

If your Python implementation is correct, your code must return: 
* 43 - which is the answer to the maximum subarray problem, and 
* <7, 10> -the start and the end indices of the max subarray. 



<!-- grade_id: cell-e4a632ce0edc1697, points: 0, grade: true, solution: true -->
```python
A = [13, -3, -25, 20, -3, -16, -23, 18, 20, -7, 12, -5, -22, 15, -4, 7]
bruteforce_max_subarray(A)
```

Output:
```
(7, 10, 43)
```

## Question 5. Asymptotic notation [time estimate: 5 minutes]
Complete the following table using the asymptotic notation that best describes the problem. For example, if both $O(n^3)$ and $O(n)$ are possible for an algorithm, the answer is $O(n)$ because the function $f(n) = O(n)$ provides a tighter and more accurate fit; if both $O(n)$ and $\Theta(n)$ are possible, the correct answer is $\Theta(n)$ because $\Theta(n)$ provides both information about the upper and lower bound, thus it is more accurate than $O(n)$.

You should copy the following table and paste and edit it in the cell below. 

Algorithm | Running time 
--- | --- 
Insertion sort | 
Selection sort |  
Bubble sort |  
Finding maximum subarray |  

<!-- grade_id: cell-247c51df276b622e, points: 0, grade: true, solution: true -->
Algorithm | Running time 
--- | --- 
Insertion sort | Θ(𝑛^2) 
Selection sort | Θ(𝑛^2) 
Bubble sort |  Θ(𝑛^2) 
Finding maximum subarray |  Θ(𝑛^2), Θ(𝑛^3)

## [Optional] Question 6 [time estimate: 15 minutes] 
How can you change this code to make it find the minimum-subarray?

<!-- grade_id: cell-f4553f71858d1bc4, points: 0, grade: false, solution: true -->
```python
def bruteforce_min_subarray(A):
    """
    Implements brute-force minimum subarray finding.
    
    Parameters
    ----------
    A : list
        a NON-EMPTY list of floats
    
    Returns
    -------
    tuple
        - the start index of the minimum subarray
        - the end index of the minimum subarray
        - the value of the minimum subarray
        
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-4263c7494a5f09d3, points: 0, grade: true, solution: false -->
```python
assert(bruteforce_min_subarray([1]*10)[0] == bruteforce_min_subarray([1]*10)[1])
assert(bruteforce_min_subarray([1]*10)[2] == 1)
```